A Synergistic Workspace for Human Consciousness Revealed by Integrated Information Decomposition

Author response:

The following is the authors’ response to the previous reviews.

We are pleased that Reviewer 3 has deemed our revisions satisfactory; below, we provide responses to the remaining Recommendations for the Authors from Reviewer 2.

Reviewer #2 (Recommendations For The Authors):

Minor corrections:

Line 91: GWT should be GNWT

Fixed, thank you.

Figure 2: fix the label "Participationcoefficient rank" (no space between Participation and coefficient)

Fixed, thank you for spotting.

Line 317: Figure 2 should be Figure 3

Fixed, thank you.

Line 360: Figure 4D, right?

Fixed, thank you. We also confirm that Figure 4 and its caption are correct. Under anaesthesia, many regions have more Integrated Information than during Recovery (red regions), but the only changes that are consistently observed across all three contrasts are the decreases.

Line 375: Should be Figure 5A

Fixed, thank you.

The recovery period of the anesthesia data is not described in Methods.

We have now added the missing information:

“Propofol was discontinued following the deep anaesthesia scan, and participants reached level 2 of the Ramsey scale approximately 11 minutes afterwards, as indicated by clear and rapid responses to verbal commands. This corresponds to the “recovery” period 176.”

We have also expanded our discussion on the interaction between information decomposition and measures of directionality:

“Indeed, transfer entropy can itself be decomposed into information-dynamic atoms through Partial Information Decomposition and Integrated Information Decomposition 33,34,49,151; ΦID can further decompose the Normalised Directed Transfer Entropy measure used by Deco et al 5, as recently demonstrated 152. We look forward to a more refined conceptualization of the synergistic workspace architecture that takes into account both information types and the directionality of information flow – especially in datasets with higher temporal resolution.”

Changes from the earlier version to the revised version

Abstract

Abstract

A central goal of neuroscience is to understand how the brain orchestrates information from multiple input streams into a unified conscious experience. Here, we address two fundamental questions: how is the human information-processing architecture functionally organised, and how does its organisation support consciousness? We combine network science and a rigorous information-theoretic notion of synergy to delineate a “synergistic global workspace”, comprising gateway regions that gather synergistic information from specialised modules across the brain. This information is then integrated within the workspace and widely distributed via broadcaster regions. Through functional MRI analysis, we show that gateway regions of the synergistic workspace correspond to the brain’s default mode network, whereas broadcasters coincide with the executive control network. Demonstrating the empirical relevance of our proposed architecture for neural information processing, we show that loss of consciousness due to general anaesthesia or disorders of consciousness corresponds to a diminished ability of the synergistic workspace to integrate information, which is restored upon recovery. Thus, loss of consciousness coincides with a breakdown of information integration within the synergistic workspace of the human brain. This work contributes to conceptual and empirical reconciliation between two prominent scientific theories of consciousness, the Global Neuronal Workspace and Integrated Information Theory. Taken together, this work provides a new perspective on the role of prominent resting-state networks within the human information-processing architecture, while also advancing our understanding of how the human brain supports consciousness through the synergistic integration of information.

Introduction

Humans and other vertebrates rely on a centralised nervous system to process information from the environment, obtained from a wide array of sensory sources. Information from different sensory sources must eventually be combined - and integrated - with the organism’s memories and goals, in order to guide adaptive behaviour effectively 1. However, understanding how the brain’s information-processing architecture enables the integration of information remains a key open challenge in neuroscience 2,3. Theoretical and empirical work in cognitive neuroscience indicates that information processed in parallel by domain-specific sensory modules needs to be integrated within a multimodal “central executive” 4. Indeed, recent work has identified subsets of regions that are consistently recruited across a variety of tasks 5–7, situated at the convergence of multiple anatomical, functional, and neurochemical hierarchies in the brain 8–20.

Prominent theories in cognitive and computational neuroscience have also proposed that global integration of information from diverse sources plays a fundamental role in relation to human consciousness 21,22. The influential Global Neuronal Workspace Theory (GNWT) focuses on the process by which specific neural information becomes available for conscious access, as occurring through the global integration induced by a “global workspace” 23–26. Within the workspace, relevant information from different sources is integrated and subsequently broadcasted back to the entire brain, in order to inform further processing and achieve “experiential integration” of distributed cortical modules into a coherent whole 24,25,27. Thus, the global workspace is attributed both the role of integrator, and the role of orchestrator of cognitive function. Also highlighting the importance of integration, the prominent Integrated Information Theory (IIT) 21,28,29 posits that the degree of consciousness in a system is determined by its “integrated information”: the amount of intrinsic information generated by the dynamics of the system considered as a whole, over and above the information generated by the dynamics of its individual constituent parts 21,28–30. Thus, this notion of integrated information corresponds to the extent to which “the whole is greater than the sum of its parts” 31.

Earlier version: Therefore, leading theoretical accounts of consciousness converge on this point: consciousness critically depends on the capability for global integration across a network of differentiated modules. Despite agreeing on the fundamental importance of information integration 32, these theories differ on its specific role and corresponding neural mechanisms. In contrast to GNWT’s account, whereby integration is viewed as a necessary - but not sufficient - prerequisite step on the way to broadcasting and consciousness, IIT proposes a more fundamental identity between consciousness and the integration of information, but without specifying a formal architecture for this process: that is, according to IIT any system that integrates information will thereby be conscious, regardless of its specific organisation31. Seen under this light, it becomes apparent that IIT and GWT are actually addressing different aspects of consciousness, and their views of integration are different but potentially complementary.

Revised version: Therefore, leading theoretical accounts of consciousness converge on this point: consciousness critically depends on the capability for global integration across a network of differentiated modules. Despite agreeing on the fundamental importance of information integration 32, these theories differ on its specific role and corresponding neural mechanisms. In contrast to GNWT’s account, whereby integration is viewed as a necessary - but not sufficient - prerequisite step on the way to broadcasting and consciousness, IIT proposes a more fundamental identity between consciousness and the integration of information, but without specifying a formal architecture for this process: that is, according to IIT any system that integrates information will thereby be conscious, regardless of its specific organisation 31. Seen under this light, it becomes apparent that IIT and GNWT are actually addressing different aspects of consciousness, and their views of integration are different but potentially complementary.

Crucially, our ability to make sense of any information-processing architecture is limited by our understanding of the information that is being processed. An elegant formal account of information in distributed systems - such as the human brain - is provided by the framework of Partial Information Decomposition (PID) 33 which extends the formalism of Shannon mutual information by demonstrating that not all information is equal. Mutual information quantifies the reduction in uncertainty about one variable, when another variable is taken into account. In the case when more than one source of information is present, PID demonstrates that two sources can possess information about a given target that is unique (each source provides independent information), redundant (the same information is provided by both sources) or synergistic (complementary information, a higher-order kind of information that is available only when both sources are considered together). As an example, humans have two sources of visual information about the world: two eyes. The information that is lost when one eye is closed is called the “unique information” of that source - information that cannot be obtained from the remaining eye. The information that one still has when one eye is closed is called “redundant information” - because it is information that is carried equally by both sources. This provides robustness: you can still see even after losing one eye. However, losing one eye also deprives you of stereoscopic information about depth. This information does not come from either eye alone: you need both, in order to perceive the third dimension. Therefore, this is called the “synergistic information” between the sources - the extra advantage that is derived from combining them. Synergistic information therefore reflects the meaning of integration-as-cooperation, whereby elements are distinct from each other, but complementary 34.

Adding to the rich literature that addresses neural information from the perspective of encoding and decoding of task variables 35, there is growing appreciation that distinct types of information – as identified by information decomposition - may play a key role in the distributed information-processing architecture of the brain 34,36–42,42,43. Information decomposition can be applied to neural data from different scales, from electrophysiology to functional MRI, with or without reference to behaviour 34. When behavioural data are taken into account, information decomposition can shed light on the processing of “extrinsic” information, understood as the translation of sensory signals into behavioural choices across neurons or regions 41,43–45. However, information decomposition can also be applied to investigate the “intrinsic” information that is present in the brain’s spontaneous dynamics in the absence of any tasks, in the same vein as resting-state “functional connectivity” and methods from statistical causal inference such as Granger causality 46. In this context, information processing should be understood in terms of the dynamics of information: where and how information is stored, transferred, and modified 34. Specifically, since the future state of the brain is at least in part determined by its previous state, it is possible to view the future state of neural units (be they regions or neurons) as the target, and ask how it is determined by the same units’ previous state, and the previous state of other units, which become the sources of information. Then, redundancy between two units occurs when their future spontaneous evolution is predicted equally well by the past of either unit. Synergy instead occurs when considering the two units together increases the mutual information between the units’ past and their future – suggesting that the future of each is shaped by its interactions with the other. At the microscale (e.g., for spiking neurons) this phenomenon has been suggested as reflecting “information modification” 36,40,47. Synergy can also be viewed as reflecting the joint contribution of parts of the system to the whole, that is not driven by common input48.

By applying a recent generalisation of PID for timeseries data – known as Integrated Information Decomposition 42,49 – we developed an information-resolved approach to decompose the information carried by brain dynamics and their intrinsic fluctuations 50. Traditional measures of statistical association (“functional connectivity”) cannot disentangle synergy and redundancy; in fact, recent work has demonstrated that functional connectivity predominantly reflects redundant interactions 34,50,51. In contrast, applying our information-resolved framework to functional MRI recordings of the human brain revealed that different regions of the human brain predominantly rely on different kinds of information for their interactions with other regions. Through this approach, we identified a “synergistic core” of brain regions supporting higher-level cognitive functions in the human brain through the synergistic integration of information 50. Similar results of a synergistic architecture were recently and independently obtained using a different decomposition (based on entropy rather than mutual information) 51.

We also observed that a synergy-based measure of emergent dynamics in functional MRI recordings is disrupted in patients suffering from chronic disorders of consciousness 52. Building on these findings, it is natural to ask whether this synergistic core could correspond to the brain’s global workspace. Furthermore, given that the views on information integration put forward by GNWT and IIT are potentially complementary, an important challenge to move the field forward is to leverage both accounts into a unified architecture that could explain empirical effects observed in neuroimaging data.

Therefore, this work sets out to address two fundamental questions of contemporary neuroscience:

How is the cognitive architecture of the human brain functionally organised, from an information-theoretic standpoint? Specifically, what brain regions does it involve, and what are the roles of the two kinds of information integration proposed by GNWT and IIT within this architecture?

How are different types of information in the brain related to human consciousness?

To address these questions, and provide an information-resolved view of human consciousness, here we study three resting-state fMRI datasets: (i) N=100 subjects from the Human Connectome Project; (ii) N=15 healthy volunteers who were scanned before and after general anaesthesia with the intravenous propofol as well as during post-anaesthetic recovery 53; (iii) N=22 patients suffering from chronic disorders of consciousness (DOC) as a result of severe brain injury 53. By comparing functional brain scans from transient anaesthetic-induced unconsciousness and from the persistent unconsciousness of DOC patients, which arises from brain injury, we can search for common brain changes associated with loss of consciousness – thereby disambiguating what is specific to loss of consciousness.

Results

Adopting an information-resolved view, we propose to divide the information-processing stream within the human brain in three key stages: (i) gathering of information from multiple distinct modules into a workspace; (ii) integration of the gathered information within the workspace; and (iii) global information broadcasting to the rest of the brain. Furthermore, we propose that while all workspace regions are involved in stage (ii), they are differentially involved in stages (i) and (iii).

The existence of a synergistic workspace and these three processing stages can be seen as emerging from a trade-off between performance and robustness that is inherent to distributed systems. Theoretical work in cognitive science 26 and the field of distributed signal processing 54,55 has long recognised the computational benefits of combining multiple distinct processing streams. However, having a single source of inputs to and outputs from the workspace introduces what is known as a “single point of failure,” which can lead to catastrophic failure in case of damage or malfunction 56. Therefore, a natural solution is to have not a single but multiple units dedicated to gathering and broadcasting information, respectively, thereby forming a workspace that can be in charge of synthesising the results of peripheral processing 57.

Pertaining to Stage (ii), we previously identified which regions of the human brain predominantly entertain synergistic interactions, and thus are most reliant on combining information from other brain regions 50 (Figure S1). The key signature of workspace regions is to have a high prevalence of synergistic (compared to redundant) functional interactions, and therefore the synergy-rich regions that we discovered are ideally poised as GNW candidates. Here, we consider the architecture of the global workspace more broadly, and combine Integrated Information Decomposition with graph-theoretical principles to bring insights about processing stages (i) and (iii) (Figure 1). We term this proposal the “Synergy-Φ-Redundancy” neurocognitive architecture (SAPHIRE) (Figure 1).

Schematic of the proposed SAPHIRE neurocognitive architecture.

Below, specialised modules characterised by robust redundant functional interactions process information about the environment. Information is then collected by workspace gateways through synergistic interactions [Stage (i)]; synergistic interactions integrate information within the synergistic global workspace [Stage (ii)]; workspace broadcasters spread the integrated information back to the specialised modules, through redundant interactions [Stage (iii)], for further processing and to guide behaviour. Orange links represent redundant interactions, and violet links represent synergistic interactions. Grey arrows represent interactions between the system and its environment.

We note that brain regions through which information gains access to the workspace should exhibit synergistic functional interactions that are widely distributed across the brain, as - by definition - the workspace gathers and synthesises information from a multiplicity of diverse brain modules. Thus, we postulate that regions that mediate the access to the synergistic workspace are functionally connected with multiple modules within networks of synergistic interactions, synthesising incoming inputs from diverse sources 58,59. We refer to such regions as gateways (Figure 1, violet nodes). In contrast, the process of broadcasting information corresponds to disseminating multiple copies of the same information from the workspace to many functionally adjacent brain regions. Therefore, broadcaster regions also have functional interactions with many different modules, but of non-synergistic, redundant interactions: “redundancy” accounts for the fact that multiple copies of the same information are being distributed. These regions are designated as broadcasters (Figure 1, orange nodes).

One approach to operationalise these ideas is by leveraging well-established graph-theoretical tools. Here, we propose to assess the diversity of intermodular functional connections using the participation coefficient 60 which captures to what extent a given node connects to many modules beyond its own (Materials and Methods). Note that this is different from the node strength, which captures a region’s total amount of connectivity, and which we used to identify which regions belong to the synergistic workspace (see Materials and Methods and Ref. 50); the participation coefficient instead quantifies the diversity of modules that a region is connected to. Therefore, gateways are identified from rs-fMRI data as brain regions that (a) belong to the workspace (i.e., have high total synergy), and (b) have a highly-ranked participation coefficient in terms of synergistic functional interactions. Conversely, broadcasters are global workspace regions (i.e., also having high synergy) that have a highly-ranked participation coefficient rank for redundant interactions.

In other words, we identify the synergistic workspace as regions where synergy predominates, which as our previous research has shown, are also involved with high-level cognitive functions and anatomically coincide with transmodal association cortices at the confluence of multiple information streams 50. This is what we should expect of a global workspace. Subsequently, to discern broadcasters from gateways within the synergistic workspace, we seek to encapsulate the meaning of a “broadcaster” in information terms. We argue that this corresponds with making the same information available to multiple modules. Sameness of information corresponds to redundancy, and connection with multiple modules can be reflected in the network-theoretic notion of participation coefficient. Thus, a broadcaster is a region in the synergistic workspace (i.e., a region with strong synergistic interactions) that in addition has a high participation coefficient for its redundant interactions.

To explore these hypotheses, we quantified synergistic and redundant interactions between 454 cortical and subcortical brain regions 61,62 based on resting-state functional MRI data from 100 subjects of the Human Connectome Project 50. Specifically, we systematically applied Integrated Information Decomposition to groups of four variables: the past and future of region X, and the past and future of region Y, for all combinations of X and Y. This provided us with a full decomposition of how information is jointly conveyed by X and Y (redundantly, uniquely, or synergistically) across time. In particular, following our previous work 50 we focused on the persistent synergy (henceforth simply synergy) and persistent redundancy (henceforth simply redundancy), which correspond to the information that is always carried synergistically (respectively, redundantly) by X and Y.

We then subdivided the brain into the well-established resting-state networks identified by Yeo and colleagues 63 plus an additional subcortical module 64. Based on this partition into modules, we identified gateways and broadcasters by comparing the participation coefficients of synergistic versus redundant interactions, for brain regions belonging to the synergistic workspace previously identified (we show a significant correlation for participation coefficient obtained from modules defined a priori as the well-known resting-state networks, or defined in a data-driven fashion from Louvain community detection 65 (Figure S2).

Intriguingly, our results reveal that gateways reside primarily in the brain’s default mode network (Figure 2B, violet). In contrast, broadcasters are mainly located in the executive control network, especially lateral prefrontal cortex (Figure 2B, orange). Remarkably, the latter results are in line with Global Neuronal Workspace Theory, which consistently identifies lateral prefrontal cortex as a major broadcaster of information 24,66.

Gateways and broadcaster regions identified by their network connectivity profiles.

(A) Group-average matrix of synergistic interactions between regions of the 454-ROI augmented Schaefer atlas. (B) Group-average matrix of redundant interactions. For ease of visualization, the colorscale in (B) pertains to log-transformed values. We highlighted modular allegiance to the canonical resting-state networks by using the colour scheme shown in between A and B. (C) Regions are identified as gateways (violet) or broadcasters (orange) based on the difference between rank of participation coefficient for synergy and redundancy, (only shown for brain regions identified as belonging to the synergistic global workspace, as per Luppi et al., 2022). Violet indicates synergy rank > redundancy rank, corresponding to workspace regions that combine information of many brain modules (gateways); orange indicates the opposite, identifying workspace regions that broadcast information to many modules. Inset: illustration of the synergistic workspace. Legend: DMN, default mode network. Som, somatomotor network. Vis, visual network. VAN, ventral attention network. DAN, dorsal attention network. FPN, fronto-parietal control network. Lim, limbic network. Sub, subcortical network (comprised of 54 regions of the Tian 2020 atlas 64). These results were also replicated using an alternative parcellation with 232 cortical and subcortical nodes (Figure S3).

Information decomposition identifies a synergistic core supporting human consciousness

Having introduced a taxonomy within the synergistic global workspace based on the distinct informational roles of different brain regions, we then sought to investigate their role in supporting human consciousness. Given the importance attributed to integration of information by both GNWT and IIT, we expected to observe reductions in integrated information within the areas of the synergistic workspace associated with loss of consciousness. Furthermore, we also reasoned that any brain regions that are specifically involved in supporting consciousness should “track” the presence of consciousness: the reductions should occur regardless of how loss of consciousness came about, and they should be restored when consciousness is regained.

Earlier version: We tested these hypotheses with resting-state fMRI from 15 healthy volunteers who were scanned before, during, and after anaesthesia with the intravenous agent propofol, as well as 22 patients with chronic disorders of consciousness (DOC) 53. Resting-state fMRI data were parcellated into 400 cortical regions from the Schaefer atlas, and 54 subcortical brain regions from the Tian atlas (same parcellation as for the previous analysis). Building on the IIT literature, which provides a formal definition of integrated information, we assessed integration corresponding to conscious activity via two alternative metrics: the well-known whole-minus-sum Φ measure introduced by Balduzzi and Tononi 31, and the “revised Φ” (ΦR) measure recently introduced by Mediano, Rosas and colleagues 49 (Materials and Methods and Figure 2). Being demonstrably non-negative, this revised measure overcomes a major conceptual limitation of the original formulation of integrated information 49.

Revised version: We tested these hypotheses with resting-state fMRI from 15 healthy volunteers who were scanned before, during, and after anaesthesia with the intravenous agent propofol, as well as 22 patients with chronic disorders of consciousness (DOC) 53. Resting-state fMRI data were parcellated into 400 cortical regions from the Schaefer atlas, and 54 subcortical brain regions from the Tian atlas (same parcellation as for the previous analysis). Building on the IIT literature, which provides a formal definition of integrated information, we assessed integration corresponding to conscious activity via two alternative metrics: the well-known whole-minus-sum Φ measure introduced by Balduzzi and Tononi 31, and the “revised Φ” (ΦR) measure recently introduced by Mediano, Rosas and colleagues 49 (Materials and Methods and Figure 3). Being demonstrably non-negative, this revised measure overcomes a major conceptual limitation of the original formulation of integrated information 49.

Added in the revised version: Integrated information decomposition.

Added in the revised version: Integrated information decomposition identifies how two sources X and Y jointly convey information across time, corresponding to all possible 4×4 combinations of redundancy, unique information (of X and of Y), and synergistic information. This decomposition highlights why the original whole-minus-sum Φ measure introduced by Balduzzi and Tononi can be negative: because it involves the subtraction of the persistent redundancy that is present in the system, leading to negative values in systems that are redundancy-dominated. This shortcoming can be corrected with the revised measure of Φ, termed ΦR.

For each subject, we computed the integrated information between each pair of BOLD signal timeseries, resulting in a 454-by-454 matrix of integrated information between brain regions. Treating this matrix as an (undirected) network enabled us to study consciousness-related changes in integrated information across conditions, which were analysed using the Network Based Statistic correction for multiple comparisons 67. Importantly, since we are interested in changes that are shared between the DOC and propofol datasets, we computed edge-level statistics using a composite null hypothesis test designed to detect such shared effects (Materials and Methods).

Analysis based on ΦR revealed a widespread reorganisation of integrated information throughout the brain when comparing awake volunteers against DOC patients, with both increases and decreases being observed (p < 0.001; Figure 4A). Likewise, propofol anaesthesia was also characterised by significant changes in integrated information between brain regions, both when compared with pre-anaesthetic wakefulness (p < 0.001; Figure 4B) and post-anaesthetic recovery (p < 0.001; Figure 4C).

Removed in the revised version: Our analysis identified a number of the ΦR connections that were reduced when consciousness was lost due to both anaesthesia and brain injury, and were restored during post-anaesthetic recovery - as we had hypothesised (Figure 4D). Remarkably, almost all regions showing consistent decreases in ΦR when consciousness was lost were members of the global synergistic workspace, and specifically located in the default mode network (bilateral precuneus and medial prefrontal cortex) - and bilateral inferior parietal cortex – although left temporal cortices were also involved (Figure 4D). Additionally, some connections exhibited increases in ΦR during loss of consciousness, and were restored upon recovery (Figure 3=4D), including areas in frontal cortex - especially lateral prefrontal cortex. Nevertheless, the overall balance was in favour of reduced integrated information: sum of F-scores associated with significant edges = −25.37 (Figure S4).

Removed in the revised version: Integrated information decomposition

Removed in the revised version: Integrated information decomposition identifies how two sources X and Y jointly convey information across time, corresponding to all possible 4×4 combinations of redundancy, unique information (of X and of Y), and synergistic information. This decomposition highlights why the original whole-minus-sum Φ measure introduced by Balduzzi and Tononi can be negative: because it involves the subtraction of the persistent redundancy that is present in the system, leading to negative values in systems that are redundancy-dominated. This shortcoming can be corrected with the revised measure of Φ, termed ΦR.

Loss of consciousness induces similar reorganisation of cortical integrated information across anaesthesia and disorders of consciousness.

Top: Brain regions exhibiting overall NBS-corrected increases (red) and decreases (blue) in integrated information exchange when consciousness is lost. (A) DOC patients minus awake healthy volunteers; (B), propofol anaesthesia minus pre-induction wakefulness; (C) propofol-anaesthesia minus post-anaesthetic recovery. (D) Overlaps between the three contrasts in (A-C), showing increases and decreases that are common across anaesthesia and disorders of consciousness.

Added in the revised version: Our analysis identified a number of the ΦR connections that were reduced when consciousness was lost due to both anaesthesia and brain injury, and were restored during post-anaesthetic recovery - as we had hypothesised (Figure 4D). Remarkably, almost all regions showing consistent decreases in ΦR when consciousness was lost were members of the global synergistic workspace, and specifically located in the default mode network (bilateral precuneus and medial prefrontal cortex) - and bilateral inferior parietal cortex – although left temporal cortices were also involved (Figure 4D). Additionally, some connections exhibited increases in ΦR during loss of consciousness, and were restored upon recovery (Figure 4D), including areas in frontal cortex - especially lateral prefrontal cortex. Nevertheless, the overall balance was in favour of reduced integrated information: sum of F-scores associated with significant edges = −25.37 (Figure S4).

These results were in contrast with the analysis based on the original formulation of Φ introduced by Balduzzi and Tononi 31, which did not identify any reductions in integrated information that were common across anaesthesia and disorders of consciousness, instead only identifying common increases (Figure S5A).

Having identified the subset of brain regions that are reliably associated with supporting human consciousness in terms of their integrated information, the last step of our analysis was to leverage the architecture proposed above to understand their role in our information-based view of the global workspace. Since IIT predicts that loss of consciousness corresponds to reductions in integrated information, we focused on regions exhibiting reliable reductions in ΦR when consciousness is lost (whether due to anaesthesia or DOC), which were restored upon recovery (shown in blue in Figure 4D).

Earlier version: Remarkably, our whole-brain results show that ΦR disconnections induced by loss of consciousness play the role of gateway nodes (Figure 4A, violet) rather than broadcaster nodes (Figure 5A, orange) according to our previous identification of gateways and broadcasters from the Human Connectome Project dataset (see Figure 2B, violet regions). Indeed, all reductions occur specifically within the default mode network (Figure 5B). Thus, these results suggest that loss of consciousness across anaesthesia and disorders of consciousness would correspond to anterior-posterior disconnection - in terms of integrated information - between DMN nodes that act as gateways into the synergistic workspace.

Revised version: Remarkably, our whole-brain results show that ΦR disconnections induced by loss of consciousness play the role of gateway nodes (Figure 5A, violet) rather than broadcaster nodes (Figure 5A, orange) according to our previous identification of gateways and broadcasters from the Human Connectome Project dataset (see Figure 2B, violet regions). Indeed, all reductions occur specifically within the default mode network (Figure 5B). Thus, these results suggest that loss of consciousness across anaesthesia and disorders of consciousness would correspond to anterior-posterior disconnection - in terms of integrated information - between DMN nodes that act as gateways into the synergistic workspace.

Synergistic core of human consciousness.

(A) Surface representations of medial and lateral views of the brain (L indicates left, R indicates right). Colours indicate brain regions that belong to the synergistic workspace, as identified from HCP data. Orange indicates broadcasters, and violet indicates gateways. Blue indicates regions that exhibit an overall significant reduction in integrated information across anaesthesia and disorders of consciousness. All blue regions overlap with violet ones. The network representation indicates edges with significantly reduced integrated information (ΦR) during both propofol anaesthesia and disorders of consciousness. The color of the nodes (violet or orange) indicates whether the corresponding regions are workspace gateways (violet) or broadcasters (orange); all regions are gateways (violet). (B) Circular graph representation of significant reductions in integrated information (ΦR) between brain regions, observed in all three contrasts, indicating membership of canonical resting-state networks. Connections indicate pairs of regions with a significant decrease of integrated information. Colour of the circle border indicates the RSN affiliation of the corresponding regions. Legend: DMN, default mode network. Som, somatomotor network. Vis, visual network. VAN, ventral attention network. DAN, dorsal attention network. FPN, fronto-parietal control network. Lim, limbic network. Sub, subcortical network (comprised of 54 regions of the Tian 2020 atlas).

Robustness and sensitivity analysis

To ensure the robustness of our results to analytic choices, we also replicated them using an alternative cortical parcellation of lower dimensionality: we used the Schaefer scale-200 cortical parcellation 62 complemented with the scale-32 subcortical ROIs from the Tian subcortical atlas 64 (Figure S5B). Additionally, we also show that our results are not dependent on the choice of parameters in the NBS analysis, and are replicated using an alternative threshold definition for the connected component (extent rather than intensity) or a more stringent value for the cluster threshold (F > 12) (Figure S5C-D). Importantly, whereas the increases in ΦR are not the same across different analytic approaches, reductions of ΦR in medial prefrontal and posterior cingulate/precuneus are reliably observed, attesting to their robustness.

Discussion

Architecture of the synergistic global workspace

This paper proposes an informational perspective on the brain’s functional architecture at the macroscale, which leverages insights from network science and a refined understanding of neural information exchange. The synergy-Φ-redundancy (SAPHIRE) architecture posits the existence of a “synergistic workspace” of brain regions characterised by highly synergistic global interactions, which we previously showed to be composed by prefrontal and parietal cortices that are critical for higher cognitive functions 50. This workspace is further functionally decomposed by distinguishing gateways, which bring information from localised modules into the workspace, and broadcasters, which disseminate multiple copies of workspace information back to low-level regions.

Remarkably, our results on the HCP dataset show that the proposed operationalisation of gateways and broadcasters corresponds to the distinction between the brain’s default mode network and executive control network, respectively. This data-driven identification of workspace gateways and broadcasters with the DMN and FPN provides a new framework to explain well-known functional differences between DMN and FPN, based on their distinct and complementary roles within the brain’s synergistic global workspace, which is discussed below.

The fronto-parietal executive control network (FPN) mainly comprises lateral prefrontal and parietal cortices, and it is associated with performance of a variety of complex, cognitively demanding tasks 68–70. A key component of this network is lateral prefrontal cortex (LPFC). Based on theoretical and empirical evidence, as summarised in a recent review of GNWT 24, this region is posited to play a major role in the global workspace, as a global broadcaster of information. Remarkably, this is precisely the role that our results assigned to LPFC, based on its combined information-theoretic and network properties. These results are also consistent with recent insights from network neuroscience, which indicate that the FPN is ideally poised to steer whole-brain dynamics through novel trajectories, in response to complex task demands 70,71. Specifically, by broadcasting to the rest of the brain information that has been integrated within the workspace, the FPN may act as global coordinator of subsequent whole-brain dynamics.

On the other hand, the default mode network comprises posterior cingulate and precuneus, medial prefrontal cortex, and inferior parietal cortices 72–74. This network, whose constituent regions have undergone substantial developments in the course of human evolution 75,76, was found to occupy a crucial position at the convergence of functional gradients of macroscale cortical organization 15,77,78, forming a structural and functional core of the human brain 79–81, in line with its recently observed involvement in cognitive tasks 82–84. In particular, the DMN is prominently involved in self-referential processing 85,86, and ‘mental-time-travel’ 87 or episodic memory and future-oriented cognition 88–91. Its posterior regions in particular, act as relays between the neocortex and the hippocampal memory system 89. Thus, in terms of both neuroanatomical connectivity and functional engagement, the DMN is uniquely positioned to integrate and contextualise information coming into the synergistic global workspace (e.g. from sensory streams) by combining it with rich information pertaining to one’s past experiences and high-level mental models about ‘self’ and world 78,92–96 - coinciding with the results of the present analysis, which identify DMN nodes as gateways of inputs to the synergistic global workspace.

It is worth noting that the role of the FPN-DMN tandem in supporting consciousness has been suggested by Shanahan’s hypothesis of a ‘connective core’ along the brain’s medial axis 97. While Shanahan’s hypotheses were primarily based on structure, in this work we combine novel information-theoretic tools to confirm and expand the connective core hypothesis from a functional, information-centric perspective, in a way that differentiates the multiple roles played by the different regions that together comprise this connective core (Figures 2,5).

Integrated Information Decomposition of human consciousness

After identifying the neuroanatomical-functional mapping of the synergistic workspace in terms of gateways and broadcasters, we sought to identify their role in supporting human consciousness. Considering integrated information as a marker of consciousness (without necessarily assuming the two to be identical), we focused on identifying regions where information integration is reduced when consciousness is lost (regardless of its cause, be it propofol anaesthesia or severe brain injury), and restored upon its recovery. Our results indicate that brain regions exhibiting consciousness-specific reductions in integrated information coincide with major nodes of the synergistic global workspace.

Intriguingly, we found that the main disruptions of information integration were localised in gateway nodes, rather than broadcasters. Thus, loss of consciousness in both anaesthesia and disorders of consciousness could be understood as a breakdown of the entry points to the “synergistic core” (Figure 5), which becomes unable to properly integrate inputs for the workspace. Importantly, the original “whole-minus-sum” Φ introduced by Balduzzi and Tononi 31 did not show consistent reductions during loss of consciousness. Thus, the present results demonstrate the empirical validity of the “revised” measure, ΦR, in addition to its theoretical soundness 49. Since workspace gateway regions coincide with the brain’s default mode network, these results are also in line with recent evidence that information content and integrative capacity of the DMN are compromised during loss of consciousness induced by both anaesthesia and severe brain injury 53,98–106, and even COVID-19 107. Due to its prominent role in self-referential processing 86, breakdown of DMN connectivity within the synergistic workspace may be seen as a failure to integrate one’s self-narrative into the “stream of consciousness”, in the words of William James.

This notion is further supported by focusing on reductions of integrated information during anaesthesia compared with wakefulness. In addition to the synergistic core, overall reductions are also observed in a set of thalamic, auditory and somatomotor regions, largely resembling the brain regions that stop responding to sensory (auditory and noxious) stimuli once the brain reaches propofol-induced saturation of EEG slow-wave activity (SWAS 108). Although there was no EEG data available to confirm this, the doses of propofol employed in the present study are compatible with the doses of propofol at which SWAS has been shown to arise 109, and therefore it is plausible that our participants also reached SWAS and the loss of brain responsiveness it indicates. Thus, both resting-state integration of information between brain regions, as well as stimulus-evoked responses within each region 108, converge to indicate that propofol disrupts further processing of thalamocortical sensory information – a phenomenon termed “thalamocortical isolation” 108. We propose that as the thalamus and sensory cortices lose their ability to respond to stimuli, they cease to provide information to the synergistic core of the global workspace, resulting in a disconnection from the external world and presumably loss of consciousness.

These results testify to the power of the Integrated Information Decomposition framework: by identifying the information-theoretic components of integrated information, we have been able to obtain insights about human consciousness that remained elusive with alternative formulations, and could not be captured via standard functional connectivity or related methods. Thus, our findings are consistent with the notion that the global workspace is relevant for supporting consciousness in the human brain, in line with the proposal that “[…] unconsciousness is not necessarily a complete suppression of information processing but rather a network dysfunction that could create inhospitable conditions for global information exchange and broadcasting” 24. GNWT postulates a key role for the global workspace in supporting consciousness: consistent with this theory, we find that several nodes of the synergistic global workspace become disconnected from each other in terms of integrated information when consciousness is lost, especially between anterior and posterior regions (Figure 4, brain networks). Thus, these are brain regions that (i) belong to the synergistic global workspace; (ii) exhibit overall reductions of integrated information when consciousness is lost; and (iii) are disconnected from other regions of the synergistic workspace when consciousness is lost. The brain regions satisfying these three conditions therefore meet the criteria for constituting an interconnected “synergistic core” of workspace regions supporting human consciousness.

Limitations and future directions

in order to obtain high spatial resolution for our identification of workspace regions, here we relied on the BOLD signal from functional MRI, which is an indirect proxy of underlying neuronal activity, with limited temporal resolution. However, we sought to alleviate potential confounds by deconvolving the hemodynamic response function from our data with a dedicated toolbox 110 (Materials and Methods), which has been previously applied both in the context of information decomposition 50, as well as anaesthetic-induced loss of consciousness 111, and disorders of consciousness 52. Additionally, the present results of an overall ΦR reduction are also broadly in line with those of a previous study 112, whose measure of synergy-redundancy balance showed, in ECoG recordings of non-human primates, a broadband shift away from synergy during anaesthesia.

Earlier version: It is also worth bearing in mind that our measure of integrated information between pairs of regions does not amount to measuring the integrated information of the brain as a whole, as formally specified in the context of Integrated Information Theory 31 - although we do show that the average integrated information between pairs of regions is overall reduced across the whole brain. We also note that our revised measure of integrated information is based on IIT 2.0 31, due to its computational tractability; as a result, it relies on a conceptually distinct understanding of integrated information from the more recent IIT 3.0 113 and IIT 4.0 114 versions. Thus, these limitations should be borne in mind when seeking to interpret the present results in the context of IIT. Indeed, future work may benefit from seeking convergence with recent advances in the characterization of emergence, which is related to integrated information 116,118,120,122.

Revised version: It is also worth bearing in mind that our measure of integrated information between pairs of regions does not amount to measuring the integrated information of the brain as a whole, as formally specified in the context of Integrated Information Theory 31 - although we do show that the average integrated information between pairs of regions is overall reduced across the whole brain. We also note that our revised measure of integrated information is based on IIT 2.0 31, due to its computational tractability; as a result, it relies on a conceptually distinct understanding of integrated information from the more recent IIT 3.0 113 and IIT 4.0 114 versions. Thus, these limitations should be borne in mind when seeking to interpret the present results in the context of IIT. Indeed, future work may benefit from seeking convergence with recent advances in the characterization of emergence, which is related to integrated information 115–118.

Earlier version: Likewise, it is not our intention to claim that ΦR and synergy, as measured at the level of regional BOLD signals, represent a direct cause of consciousness, or are identical to it. Rather, our work is intended to use these measures similarly to the use of sample entropy and Lempel-Ziv complexity for BOLD signals 53,115: as theoretically grounded macroscale indicators, whose empirical relationship to consciousness may point towards the relevant underlying neural phenomena. In other words, while our results do show that BOLD-derived ΦR tracks the loss and recovery of consciousness, we do not claim that they are the cause of it: only that an empirical relationship exists, which is in line with what we might expect on theoretical grounds. Future work will be required to identify whether this empirical relationship also holds at the microscale, and whether the causal mechanisms that induce loss of consciousness are also causally responsible for loss of integrated information.

Revised version: Likewise, it is not our intention to claim that ΦR and synergy, as measured at the level of regional BOLD signals, represent a direct cause of consciousness, or are identical to it. Rather, our work is intended to use these measures similarly to the use of sample entropy and Lempel-Ziv complexity for BOLD signals 53,119: as theoretically grounded macroscale indicators, whose empirical relationship to consciousness may point towards the relevant underlying neural phenomena. In other words, while our results do show that BOLD-derived ΦR tracks the loss and recovery of consciousness, we do not claim that they are the cause of it: only that an empirical relationship exists, which is in line with what we might expect on theoretical grounds. Future work will be required to identify whether this empirical relationship also holds at the microscale, and whether the causal mechanisms that induce loss of consciousness are also causally responsible for loss of integrated information.

Intriguingly, although we have focused on anaesthetic-induced decreases in integrated information, due to IIT’s prediction that this is what should occur during loss of consciousness, our results also indicate concomitant increases of integrated information – possibly reflecting compensatory attempts, although we refrain from further speculation (Figure 4). Interestingly, increases appear to coincide with broadcaster nodes of the synergistic workspace. In particular, even though lateral prefrontal cortices are among the regions most closely associated with the global neuronal workspace in the literature 24,66, our results indicate a paradoxical net increase in lateral prefrontal integrated information during anaesthesia and DOC. We interpret this qualitatively different behaviour as indicating that different subsets of the global workspace may be differentially involved in supporting consciousness.

Earlier version: However, we note that, whereas the decreases in integrated information were robust to the use of different analytic approaches (e.g., use of a different parcellation or different NBS threshold), the increases that we observed were less robust, with no region consistently showing increases in integrated information (Figure S5B-D). Nevertheless, both this phenomenon and the meaning of increased integrated information between brain regions deserve further investigation. Indeed, dreaming during anaesthesia has been reported to occur in up to 27% of cases 117, and behaviourally unresponsive participants have been shown to perform mental imagery tasks during anaesthesia, both of which constitute cases of disconnected consciousness 119. Thus, although our doses of propofol were consistent with the presence of SWAS, we cannot exclude that some of our participants may have been merely disconnected but still conscious, possibly driving the increases we observed.

Revised version: However, we note that, whereas the decreases in integrated information were robust to the use of different analytic approaches (e.g., use of a different parcellation or different NBS threshold), the increases that we observed were less robust, with no region consistently showing increases in integrated information (Figure S5B-D). Nevertheless, both this phenomenon and the meaning of increased integrated information between brain regions deserve further investigation. Indeed, dreaming during anaesthesia has been reported to occur in up to 27% of cases 120, and behaviourally unresponsive participants have been shown to perform mental imagery tasks during anaesthesia, both of which constitute cases of disconnected consciousness 121. Thus, although our doses of propofol were consistent with the presence of SWAS, we cannot exclude that some of our participants may have been merely disconnected but still conscious, possibly driving the increases we observed.

Earlier version: More broadly, future research may also benefit from characterising the role of the synergistic workspace in the states of altered consciousness induced e.g. by psychedelics 121,123,124, especially since prominent involvement of the DMN has already been identified 125,126. Likewise, the use of paradigms different from resting-state, such as measuring the brain’s spontaneous responses to engaging stimuli (e.g. suspenseful narratives 127 or engaging movies 128) may provide evidence for a more comprehensive understanding of brain changes during unconsciousness. Likewise, it will be of great interest to investigate whether and how reorganization of the synergistic global workspace is reflected in other indicators of consciousness 129, such as the brain’s response to external perturbations – such as the EEG response to brief magnetic pulses used to compute the Perturbational Complexity Index, one of the most discriminative indices of consciousness available to date 130–134.

Revised version: More broadly, future research may also benefit from characterising the role of the synergistic workspace in the states of altered consciousness induced e.g. by psychedelics 122–124, especially since prominent involvement of the DMN has already been identified 125,126. Likewise, the use of paradigms different from resting-state, such as measuring the brain’s spontaneous responses to engaging stimuli (e.g. suspenseful narratives 127 or engaging movies 128) may provide evidence for a more comprehensive understanding of brain changes during unconsciousness. Likewise, it will be of great interest to investigate whether and how reorganization of the synergistic global workspace is reflected in other indicators of consciousness 129, such as the brain’s response to external perturbations – such as the EEG response to brief magnetic pulses used to compute the Perturbational Complexity Index, one of the most discriminative indices of consciousness available to date 130–134.

The PCI is used as a means of assessing the brain’s current state, but stimulation protocols can also be adopted to directly induce transitions between states of consciousness. In rodents, carbachol administration to frontal cortex awakens rats from sevoflurane anaesthesia135, and optogenetic stimulation was used to identify a role of central thalamus neurons in controlling transitions between states of responsiveness136,137. Additionally, several studies in non-human primates have now shown that electrical stimulation of the central thalamus can reliably induce awakening from anaesthesia, accompanied by the reversal of electrophysiological and fMRI markers of anaesthesia 138–143. Finally, in human patients suffering from disorders of consciousness, stimulation of intra-laminar central thalamic nuclei was reported to induce behavioural improvement 144, and ultrasonic stimulation 145,146 and deep-brain stimulation are among potential therapies being considered for DOC patients 147,148. It will be of considerable interest to determine whether our corrected measure of integrated information and topography of the synergistic workspace also restored by these causal interventions.

Earlier version: Additionally, the reliance here on ‘resting-state’ data without external stimuli may have resulted in an overestimation of the DMN’s role in consciousness, and an under-estimation of the FPN (including lateral PFC), given their known different recruitment during no-task conditions 72. Indeed, recent efforts have been carried out to obtain a data-driven characterisation of the brain’s global workspace based on regions’ involvement across multiple different tasks 5. This work is complementary to ours in two aspects: first, the focus of Deco et al. 5 is on the role of the workspace related to cognition, whereas here we focus primarily on consciousness. Second, by using transfer entropy as a measure of functional connectivity, Deco and colleagues 5 assessed the directionality of information exchange – whereas our measure of integrated information is undirected, but are able to distinguish between different kinds of information being exchanged and integrated. Thus, different ways of defining and characterising a global workspace in the human brain are possible, and can provide complementary insights about distinct aspects of the human neurocognitive architecture.

Revised version: Additionally, the reliance here on ‘resting-state’ data without external stimuli may have resulted in an overestimation of the DMN’s role in consciousness, and an under-estimation of the FPN (including lateral PFC), given their known different recruitment during no-task conditions 72. Indeed, recent efforts have been carried out to obtain a data-driven characterisation of the brain’s global workspace based on regions’ involvement across multiple different tasks 5. This work is complementary to ours in two aspects: first, the focus of Deco et al. 5 is on the role of the workspace related to cognition, whereas here we focus primarily on consciousness. Second, by using transfer entropy 149,150 as a measure of functional connectivity, Deco and colleagues 5 assessed the directionality of information exchange – whereas our measure of integrated information is undirected, but are able to distinguish between different kinds of information being exchanged and integrated. Thus, different ways of defining and characterising a global workspace in the human brain are possible, and can provide complementary insights about distinct aspects of the human neurocognitive architecture. Indeed, transfer entropy can itself be decomposed into information-dynamic atoms through Partial Information Decomposition and Integrated Information Decomposition 33,34,49,151; ΦID can further decompose the Normalised Directed Transfer Entropy measure used by Deco et al 5, as recently demonstrated 152. We look forward to a more refined conceptualization of the synergistic workspace architecture that takes into account both information types and the directionality of information flow – especially in datasets with higher temporal resolution.

Earlier version: Looking forward, growing evidence indicates an important role for brain dynamics and time-resolved brain states in supporting cognition 7,18,149–156 and consciousness 53,106,157–164. Therefore, time-resolved extensions of our framework, such as developed by Varley and colleagues 42, may shed further light on the dynamics of the synergistic workspace, especially if combined with neuroimaging modalities offering higher temporal resolution, such as magneto- or electroencephalography. More broadly, a key strength of our proposed cognitive architecture is its generality: being entirely grounded in the combination of information theory and network science, it could be applied to shed light on cognition in humans and other organisms 165, but also to inspire further development of artificial cognitive systems 34,166–170.

Revised version: Looking forward, growing evidence indicates an important role for brain dynamics and time-resolved brain states in supporting cognition 7,18,153–160 and consciousness 53,106,161–168. Therefore, time-resolved extensions of our framework, such as developed by Varley and colleagues 42, may shed further light on the dynamics of the synergistic workspace, especially if combined with neuroimaging modalities offering higher temporal resolution, such as magneto-or electroencephalography. More broadly, a key strength of our proposed cognitive architecture is its generality: being entirely grounded in the combination of information theory and network science, it could be applied to shed light on cognition in humans and other organisms 169, but also to inspire further development of artificial cognitive systems 34,170–174.

Conclusion

Overall, we have shown that powerful insights about human consciousness and neurocognitive architecture can be obtained through the information-resolved approach, afforded by the framework of Integrated Information Decomposition. Importantly, the proposed criteria to identify gateways, broadcasters, and the synergistic workspace itself, are based on practical network and information-theoretic tools, which are applicable to a broad range of neuroimaging datasets and neuroscientific questions. These findings bring us closer to a unified theoretical understanding of consciousness and its neuronal underpinnings - how mind arises from matter.

Materials and Methods

Earlier version: The propofol and DOC patient functional data employed in this study have been published before 53,115,121,171–174. For clarity and consistency of reporting, where applicable we use the same wording as our previous work 53,121,171.

Revised version: The propofol and DOC patient functional data employed in this study have been published before 53,119,122,175–178. For clarity and consistency of reporting, where applicable we use the same wording as our previous work 53,122,175.

Anaesthesia Data: Recruitment

Earlier version: The propofol data were collected between May and November 2014 at the Robarts Research Institute in London, Ontario (Canada) 53. The study received ethical approval from the Health Sciences Research Ethics Board and Psychology Research Ethics Board of Western University (Ontario, Canada). Healthy volunteers (n=19) were recruited (18–40 years; 13 males). Volunteers were right-handed, native English speakers, and had no history of neurological disorders. In accordance with relevant ethical guidelines, each volunteer provided written informed consent, and received monetary compensation for their time. Due to equipment malfunction or physiological impediments to anaesthesia in the scanner, data from n=3 participants (1 male) were excluded from analyses, leaving a total n=16 for analysis 53,121,171.

Revised version: The propofol data were collected between May and November 2014 at the Robarts Research Institute in London, Ontario (Canada) 53. The study received ethical approval from the Health Sciences Research Ethics Board and Psychology Research Ethics Board of Western University (Ontario, Canada). Healthy volunteers (n=19) were recruited (18–40 years; 13 males). Volunteers were right-handed, native English speakers, and had no history of neurological disorders. In accordance with relevant ethical guidelines, each volunteer provided written informed consent, and received monetary compensation for their time. Due to equipment malfunction or physiological impediments to anaesthesia in the scanner, data from n=3 participants (1 male) were excluded from analyses, leaving a total n=16 for analysis 53,122,175.

Anaesthesia Data: Procedure

Earlier version: Resting-state fMRI data were acquired at different propofol levels: no sedation (Awake), and Deep anaesthesia (corresponding to Ramsay score of 5). As previously reported 53,121,171, for each condition fMRI acquisition began after two anaesthesiologists and one anaesthesia nurse independently assessed Ramsay level in the scanning room. The anaesthesiologists and the anaesthesia nurse could not be blinded to experimental condition, since part of their role involved determining the participants’ level of anaesthesia. Note that the Ramsay score is designed for critical care patients, and therefore participants did not receive a score during the Awake condition before propofol administration: rather, they were required to be fully awake, alert and communicating appropriately. To provide a further, independent evaluation of participants’ level of responsiveness, they were asked to perform two tasks: a test of verbal memory recall, and a computer-based auditory target-detection task. Wakefulness was also monitored using an infrared camera placed inside the scanner.

Revised version: Resting-state fMRI data were acquired at different propofol levels: no sedation (Awake), and Deep anaesthesia (corresponding to Ramsay score of 5). As previously reported 53,122,175, for each condition fMRI acquisition began after two anaesthesiologists and one anaesthesia nurse independently assessed Ramsay level in the scanning room. The anaesthesiologists and the anaesthesia nurse could not be blinded to experimental condition, since part of their role involved determining the participants’ level of anaesthesia. Note that the Ramsay score is designed for critical care patients, and therefore participants did not receive a score during the Awake condition before propofol administration: rather, they were required to be fully awake, alert and communicating appropriately. To provide a further, independent evaluation of participants’ level of responsiveness, they were asked to perform two tasks: a test of verbal memory recall, and a computer-based auditory target-detection task. Wakefulness was also monitored using an infrared camera placed inside the scanner.

Propofol was administered intravenously using an AS50 auto syringe infusion pump (Baxter Healthcare, Singapore); an effect-site/plasma steering algorithm combined with the computer-controlled infusion pump was used to achieve step-wise sedation increments, followed by manual adjustments as required to reach the desired target concentrations of propofol according to the TIVA Trainer (European Society for Intravenous Aneaesthesia, eurosiva.eu) pharmacokinetic simulation program. This software also specified the blood concentrations of propofol, following the Marsh 3-compartment model, which were used as targets for the pharmacokinetic model providing target-controlled infusion. After an initial propofol target effect-site concentration of 0.6 µg mL-1, concentration was gradually increased by increments of 0.3 µg mL1, and Ramsay score was assessed after each increment: a further increment occurred if the Ramsay score was lower than 5. The mean estimated effect-site and plasma propofol concentrations were kept stable by the pharmacokinetic model delivered via the TIVA Trainer infusion pump. Ramsay level 5 was achieved when participants stopped responding to verbal commands, were unable to engage in conversation, and were rousable only to physical stimulation. Once both anaesthesiologists and the anaesthesia nurse all agreed that Ramsay sedation level 5 had been reached, and participants stopped responding to both tasks, data acquisition was initiated. The mean estimated effect-site propofol concentration was 2.48 (1.82-3.14) µg mL-1, and the mean estimated plasma propofol concentration was 2.68 (1.92-3.44) µg mL-1. Mean total mass of propofol administered was 486.58 (373.30-599.86) mg. These values of variability are typical for the pharmacokinetics and pharmacodynamics of propofol. Oxygen was titrated to maintain SpO2 above 96%.

Earlier version: At Ramsay 5 level, participants remained capable of spontaneous cardiovascular function and ventilation. However, the sedation procedure did not take place in a hospital setting; therefore, intubation during scanning could not be used to ensure airway security during scanning. Consequently, although two anaesthesiologists closely monitored each participant, scanner time was minimised to ensure return to normal breathing following deep sedation. No state changes or movement were noted during the deep sedation scanning for any of the participants included in the study 53,121,171.

Revised version: At Ramsay 5 level, participants remained capable of spontaneous cardiovascular function and ventilation. However, the sedation procedure did not take place in a hospital setting; therefore, intubation during scanning could not be used to ensure airway security during scanning. Consequently, although two anaesthesiologists closely monitored each participant, scanner time was minimised to ensure return to normal breathing following deep sedation. No state changes or movement were noted during the deep sedation scanning for any of the participants included in the study 53,122,175. Propofol was discontinued following the deep anaesthesia scan, and participants reached level 2 of the Ramsey scale approximately 11 minutes afterwards, as indicated by clear and rapid responses to verbal commands. This corresponds to the “recovery” period 176.

Anaesthesia Data: Design

Earlier version: As previously reported 53,121,171, once in the scanner participants were instructed to relax with closed eyes, without falling asleep. Resting-state functional MRI in the absence of any tasks was acquired for 8 minutes for each participant. A further scan was also acquired during auditory presentation of a plot-driven story through headphones (5-minute long). Participants were instructed to listen while keeping their eyes closed. The present analysis focuses on the resting-state data only; the story scan data have been published separately, and will not be discussed further here.

Revised version: As previously reported 53,122,175, once in the scanner participants were instructed to relax with closed eyes, without falling asleep. Resting-state functional MRI in the absence of any tasks was acquired for 8 minutes for each participant, in each condition. A further scan was also acquired during auditory presentation of a plot-driven story through headphones (5-minute long). Participants were instructed to listen while keeping their eyes closed. The present analysis focuses on the resting-state data only; the story scan data have been published separately, and will not be discussed further here.

Anaesthesia Data: FMRI Data Acquisition

Earlier version: As previously reported 53,121,171, MRI scanning was performed using a 3-Tesla Siemens Tim Trio scanner (32-channel coil), and 256 functional volumes (echo-planar images, EPI) were collected from each participant, with the following parameters: slices = 33, with 25% inter-slice gap; resolution = 3mm isotropic; TR = 2000ms; TE = 30ms; flip angle = 75 degrees; matrix size = 64×64. The order of acquisition was interleaved, bottom-up. Anatomical scanning was also performed, acquiring a high-resolution T1-weighted volume (32-channel coil, 1mm isotropic voxel size) with a 3D MPRAGE sequence, using the following parameters: TA = 5min, TE = 4.25ms, 240×256 matrix size, 9 degrees flip angle 53,121,171.

Revised version: As previously reported 53,122,175, MRI scanning was performed using a 3-Tesla Siemens Tim Trio scanner (32-channel coil), and 256 functional volumes (echo-planar images, EPI) were collected from each participant, with the following parameters: slices = 33, with 25% inter-slice gap; resolution = 3mm isotropic; TR = 2000ms; TE = 30ms; flip angle = 75 degrees; matrix size = 64×64. The order of acquisition was interleaved, bottom-up. Anatomical scanning was also performed, acquiring a high-resolution T1-weighted volume (32-channel coil, 1mm isotropic voxel size) with a 3D MPRAGE sequence, using the following parameters: TA = 5min, TE = 4.25ms, 240×256 matrix size, 9 degrees flip angle 53,122,175.

Disorders of Consciousness Patient Data: Recruitment

Removed in the revised version: A total of 71 DOC patients were recruited from specialised long-term care centres from January 2010 to December 2015 53,121,171. Ethical approval for this study was provided by the National Research Ethics Service (National Health Service, UK; LREC reference 99/391).

Earlier version: Patients were eligible to be recruited in the study if they had a diagnosis of chronic disorder of consciousness, provided that written informed consent to participation was provided by their legal representative, and provided that the patients could be transported to Addenbrooke’s Hospital (Cambridge, UK). The exclusion criteria included any medical condition that made it unsafe for the patient to participate, according to clinical personnel blinded to the specific aims of the study; or any reason that made a patient unsuitable to enter the MRI scanner environment (e.g. non-MRI-safe implants). Patients were also excluded based on substantial pre-existing mental health problems, or insufficient fluency in the English language prior to their injury. After admission to Addenbrooke’s Hospital, each patient underwent clinical and neuroimaging testing, spending a total of five days in the hospital (including arrival and departure days). Neuroimaging scanning took place at the Wolfson Brain Imaging Centre (Addenbrooke’s Hospital, Cambridge, UK), and medication prescribed to each patient was maintained during scanning.

Revised version: A total of 71 DOC patients were recruited from specialised long-term care centres from January 2010 to December 2015 53,122,175. Ethical approval for this study was provided by the National Research Ethics Service (National Health Service, UK; LREC reference 99/391). Patients were eligible to be recruited in the study if they had a diagnosis of chronic disorder of consciousness, provided that written informed consent to participation was provided by their legal representative, and provided that the patients could be transported to Addenbrooke’s Hospital (Cambridge, UK). The exclusion criteria included any medical condition that made it unsafe for the patient to participate, according to clinical personnel blinded to the specific aims of the study; or any reason that made a patient unsuitable to enter the MRI scanner environment (e.g. non-MRI-safe implants). Patients were also excluded based on substantial pre-existing mental health problems, or insufficient fluency in the English language prior to their injury. After admission to Addenbrooke’s Hospital, each patient underwent clinical and neuroimaging testing, spending a total of five days in the hospital (including arrival and departure days). Neuroimaging scanning took place at the Wolfson Brain Imaging Centre (Addenbrooke’s Hospital, Cambridge, UK), and medication prescribed to each patient was maintained during scanning.

Earlier version: For each day of admission, Coma Recovery Scale-Revised (CRS-R) assessments were recorded at least daily. Patients whose behavioural responses were not indicative of awareness at any time, were classified as UWS. In contrast, patients were classified as being in a minimally conscious state (MCS) if they provided behavioural evidence of simple automatic motor reactions (e.g., scratching, pulling the bed sheet), visual fixation and pursuit, or localisation to noxious stimulation. Since this study focused on whole-brain properties, coverage of most of the brain was required, and we followed the same criteria as in our previous studies 53,121,171; before analysis took place, patients were systematically excluded if an expert neuroanatomist blinded to diagnosis judged that they displayed excessive focal brain damage (over one third of one hemisphere), or if brain damage led to suboptimal segmentation and normalisation, or due to excessive head motion in the MRI scanner (exceeding 3mm translation or 3 degrees rotation). Of the initial sample of 71 patients who had been recruited, a total of 22 adults (14 males; 17–70 years; mean time post injury: 13 months) meeting diagnostic criteria for Unresponsive Wakefulness Syndrome/Vegetative State or Minimally Conscious State due to brain injury were included in this study. In addition to the researcher and radiographer, a research nurse was also present during scanning. Since the patients’ status as DOC patients was evident, no researcher blinding was possible.

Revised version: For each day of admission, Coma Recovery Scale-Revised (CRS-R) assessments were recorded at least daily. Patients whose behavioural responses were not indicative of awareness at any time, were classified as UWS. In contrast, patients were classified as being in a minimally conscious state (MCS) if they provided behavioural evidence of simple automatic motor reactions (e.g., scratching, pulling the bed sheet), visual fixation and pursuit, or localisation to noxious stimulation. Since this study focused on whole-brain properties, coverage of most of the brain was required, and we followed the same criteria as in our previous studies 53,122,175; before analysis took place, patients were systematically excluded if an expert neuroanatomist blinded to diagnosis judged that they displayed excessive focal brain damage (over one third of one hemisphere), or if brain damage led to suboptimal segmentation and normalisation, or due to excessive head motion in the MRI scanner (exceeding 3mm translation or 3 degrees rotation). Of the initial sample of 71 patients who had been recruited, a total of 22 adults (14 males; 17–70 years; mean time post injury: 13 months) meeting diagnostic criteria for Unresponsive Wakefulness Syndrome/Vegetative State or Minimally Conscious State due to brain injury were included in this study. In addition to the researcher and radiographer, a research nurse was also present during scanning. Since the patients’ status as DOC patients was evident, no researcher blinding was possible.

Disorders of Consciousness Patient Data: FMRI Data Acquisition

Earlier version: As previously reported 53,121,171, resting-state fMRI was acquired for 10 minutes (300 volumes, TR=2000ms) using a Siemens Trio 3T scanner (Erlangen, Germany). Functional images (32 slices) were acquired using an echo planar sequence, with the following parameters: 3 x 3 x 3.75mm resolution, TR = 2000ms, TE = 30ms, 78 degrees FA. Anatomical scanning was also performed, acquiring high-resolution T1-weighted images with an MPRAGE sequence, using the following parameters: TR = 2300ms, TE = 2.47ms, 150 slices, resolution 1 x 1 x 1mm.

Revised version: As previously reported 53,122,175, resting-state fMRI was acquired for 10 minutes (300 volumes, TR=2000ms) using a Siemens Trio 3T scanner (Erlangen, Germany). Functional images (32 slices) were acquired using an echo planar sequence, with the following parameters: 3 x 3 x 3.75mm resolution, TR = 2000ms, TE = 30ms, 78 degrees FA. Anatomical scanning was also performed, acquiring high-resolution T1-weighted images with an MPRAGE sequence, using the following parameters: TR = 2300ms, TE = 2.47ms, 150 slices, resolution 1 x 1 x 1mm.

Functional MRI preprocessing and denoising

Earlier version: The functional imaging data were preprocessed using a standard pipeline, implemented within the SPM12-based (http://www.fil.ion.ucl.ac.uk/spm) toolbox CONN (http://www.nitrc.org/projects/conn), version 17f 175. The pipeline comprised the following steps: removal of the first five scans, to allow magnetisation to reach steady state; functional realignment and motion correction; slice-timing correction to account for differences in time of acquisition between slices; identification of outlier scans for subsequent regression by means of the quality assurance/artifact rejection software art (http://www.nitrc.org/projects/artifact_detect); structure-function coregistration using each volunteer’s high-resolution T1-weighted image; spatial normalisation to Montreal Neurological Institute (MNI-152) standard space with 2mm isotropic resampling resolution, using the segmented grey matter image, together with an a priori grey matter template.

Revised version: The functional imaging data were preprocessed using a standard pipeline, implemented within the SPM12-based (http://www.fil.ion.ucl.ac.uk/spm) toolbox CONN (http://www.nitrc.org/projects/conn), version 17f 179. The pipeline comprised the following steps: removal of the first five scans, to allow magnetisation to reach steady state; functional realignment and motion correction; slice-timing correction to account for differences in time of acquisition between slices; identification of outlier scans for subsequent regression by means of the quality assurance/artifact rejection software art (http://www.nitrc.org/projects/artifact_detect); structure-function coregistration using each volunteer’s high-resolution T1-weighted image; spatial normalisation to Montreal Neurological Institute (MNI-152) standard space with 2mm isotropic resampling resolution, using the segmented grey matter image, together with an a priori grey matter template.

Earlier version: To reduce noise due to cardiac, breathing, and motion artifacts, which are known to impact functional connectivity and network analyses 176,177, we applied the anatomical CompCor method of denoising the functional data 178, also implemented within the CONN toolbox. As for preprocessing, we followed the same denoising described in previous work 53,121,171. The anatomical CompCor method involves regressing out of the functional data the following confounding effects: the first five principal components attributable to each individual’s white matter signal, and the first five components attributable to individual cerebrospinal fluid (CSF) signal; six subject-specific realignment parameters (three translations and three rotations) as well as their first-order temporal derivatives; the artefacts identified by art; and main effect of scanning condition 178. Linear detrending was also applied, and the subject-specific denoised BOLD signal timeseries were band-pass filtered to eliminate both low-frequency drift effects and high-frequency noise, thus retaining temporal frequencies between 0.008 and 0.09 Hz.

Revised version: To reduce noise due to cardiac, breathing, and motion artifacts, which are known to impact functional connectivity and network analyses 180,181, we applied the anatomical CompCor method of denoising the functional data 182, also implemented within the CONN toolbox. As for preprocessing, we followed the same denoising described in previous work 53,122,175. The anatomical CompCor method involves regressing out of the functional data the following confounding effects: the first five principal components attributable to each individual’s white matter signal, and the first five components attributable to individual cerebrospinal fluid (CSF) signal; six subject-specific realignment parameters (three translations and three rotations) as well as their first-order temporal derivatives; the artefacts identified by art; and main effect of scanning condition 182. Linear detrending was also applied, and the subject-specific denoised BOLD signal timeseries were band-pass filtered to eliminate both low-frequency drift effects and high-frequency noise, thus retaining temporal frequencies between 0.008 and 0.09 Hz.

Earlier version: The step of global signal regression (GSR) has received substantial attention in the fMRI literature, as a potential denoising step 179–182. However, GSR mathematically mandates that approximately 50% of correlations between regions will be negative 182, thereby removing potentially meaningful differences in the proportion of anticorrelations; additionally, it has been shown across species and states of consciousness that the global signal contains information relevant for consciousness 183. Therefore, here we chose to avoid GSR in favour of the aCompCor denoising procedure, in line with previous work 53,121,171.

Revised version: The step of global signal regression (GSR) has received substantial attention in the fMRI literature, as a potential denoising step 183–186. However, GSR mathematically mandates that approximately 50% of correlations between regions will be negative 186, thereby removing potentially meaningful differences in the proportion of anticorrelations; additionally, it has been shown across species and states of consciousness that the global signal contains information relevant for consciousness 187. Therefore, here we chose to avoid GSR in favour of the aCompCor denoising procedure, in line with previous work 53,122,175.

Earlier version: Due to the presence of deformations caused by brain injury, rather than relying on automated pipelines, DOC patients’ brains were individually preprocessed using SPM12, with visual inspections after each step. Additionally, to further reduce potential movement artefacts, data underwent despiking with a hyperbolic tangent squashing function, also implemented from the CONN toolbox 175. The remaining preprocessing and denoising steps were the same as described above.

Revised version: Due to the presence of deformations caused by brain injury, rather than relying on automated pipelines, DOC patients’ brains were individually preprocessed using SPM12, with visual inspections after each step. Additionally, to further reduce potential movement artefacts, data underwent despiking with a hyperbolic tangent squashing function, also implemented from the CONN toolbox 179. The remaining preprocessing and denoising steps were the same as described above.

Brain Parcellation

Brains were parcellated into 454 cortical and subcortical regions of interest (ROIs). The 400 cortical ROIs were obtained from the scale-400 version of the recent Schaefer local-global functional parcellation 62. Since this parcellation only includes cortical regions, it was augmented with 54 subcortical ROIs from the highest resolution of the recent Tian parcellation 64. We refer to this 454-ROI parcellation as the “augmented Schaefer” 61. To ensure the robustness of our results to the choice of atlas, we also replicated them using an alternative cortical parcellation of different dimensionality: we used the Schaefer scale-200 cortical parcellation, complemented with the scale-32 subcortical ROIs from the Tian subcortical atlas 61. The timecourses of denoised BOLD signals were averaged between all voxels belonging to a given atlas-derived ROI, using the CONN toolbox. The resulting region-specific timecourses of each subject were then extracted for further analysis in MATLAB.

HRF deconvolution

In accordance with our previous work 50,52 and previous studies using of information-theoretic measures in the context of functional MRI data, we used a dedicated toolbox 110 to deconvolve the hemodynamic response function from our regional BOLD signal timeseries prior to analysis.

Measuring Integrated Information

The framework of integrated information decomposition (ΦID) unifies integrated information theory (IIT) and partial information decomposition (PID) to decompose information flow into interpretable, disjoint parts. In this section we provide a brief description of ΦID and formulae required to compute the results. For further details, see 49,50.

Partial information decomposition

Removed in the revised version: We begin with Shannon’s Mutual information (MI), which quantifies the interdependence between two random variables X and Y. It is calculated as

Earlier version: where H(X) stands for the Shannon entropy of a variable X. Above, the first equality states that the mutual information is equal to the reduction in entropy (i.e., uncertainty) about X after Y is known. Put simply, the mutual information quantifies the information that one variable provides about another 184.

Revised version: We begin with Shannon’s Mutual information (MI), which quantifies the interdependence between two random variables X and Y. It is calculated as where H(X) stands for the Shannon entropy of a variable X. Above, the first equality states that the mutual information is equal to the reduction in entropy (i.e., uncertainty) about X after Y is known. Put simply, the mutual information quantifies the information that one variable provides about another 188.

Earlier version: Crucially, Williams and Beer 33 observed that the information that two source variables X and Y give about a third target variable Z, I(X,Y; Z), should be decomposable in terms of different types of information: information provided by one source but not the other (unique information), by both sources separately (redundant information), or jointly by their combination (synergistic information). Following this intuition, they developed the Partial Information Decomposition (PID; 33) framework, which leads to the following fundamental decomposition:

Revised version: Crucially, Williams and Beer 33 observed that the information that two source variables X and Y give about a third target variable Z, I(X,Y; Z), should be decomposable in terms of different types of information: information provided by one source but not the other (unique information), by both sources separately (redundant information), or jointly by their combination (synergistic information). Following this intuition, they developed the Partial Information Decomposition (PID; 33) framework, which leads to the following fundamental decomposition: Above, Un corresponds to the unique information one source but the other doesn’t, Red is the redundancy between both sources, and Syn is their synergy: information that neither X nor Y alone can provide, but that can be obtained by considering X and Y together.

Removed in the revised version: Above, Un corresponds to the unique information one source but the other doesn’t, Red is the redundancy between both sources, and Syn is their synergy: information that neither X nor Y alone can provide, but that can be obtained by considering X and Y together.

The simplest example of a purely synergistic system is one in which X and Y are independent fair coins, and Z is determined by the exclusive-OR function Z = XOR(X,Y): i.e., Z=0 whenever X and Y have the same value, and Z=1 otherwise. It can be shown that X and Y are both statistically independent of Z, which implies that neither of them provide - by themselves - information about Z. However, X and Y together fully determine Z, hence the relationship between Z with X and Y is purely synergistic.

As another example for the case of Gaussian variables (as employed here), consider a 2-node coupled autoregressive process with two parameters: a noise correlation c and a coupling parameter a. As c increases, the system is flooded by “common noise”, making the system increasingly redundant because the common noise “swamps” the signal of each node. As a increases, each node has a stronger influence both on the other and on the system as a whole, and we expect synergy to increase. Therefore, synergy reflects the joint contribution of parts of the system to the whole that is not driven by common noise. This has been demonstrated through computational modelling 48.

Recently, Mediano et al (2021) 49 formulated an extension of PID able to decompose the information that multiple source variables have about multiple target variables. This makes PID applicable to the dynamical systems setting, and yields a decomposition with redundant, unique, and synergistic components in the past and future that can be used as a principled method to analyse information flow in neural activity (Figure 3).

Synergy and redundancy calculation

Earlier version: While there is ongoing research on the advantages of different information decompositions for discrete data, most decompositions converge into the same simple form for the case of continuous Gaussian variables 185. Known as minimum mutual information PID (MMI-PID), this decomposition quantifies redundancy in terms of the minimum mutual information of each individual source with the target; synergy, then, becomes identified with the additional information provided by the weaker source once the stronger source is known. Since linear-Gaussian models are sufficiently good descriptors of functional MRI timeseries (and more complex, non-linear models offer no advantage 191,193), here we adopt the MMI-PID decomposition, following our own and others’ previous applications of PID to neuroscientific data 50.

Revised version: While there is ongoing research on the advantages of different information decompositions for discrete data, most decompositions converge into the same simple form for the case of continuous Gaussian variables 189. Known as minimum mutual information PID (MMI-PID), this decomposition quantifies redundancy in terms of the minimum mutual information of each individual source with the target; synergy, then, becomes identified with the additional information provided by the weaker source once the stronger source is known. Since linear-Gaussian models are sufficiently good descriptors of functional MRI timeseries (and more complex, non-linear models offer no advantage 190,191), here we adopt the MMI-PID decomposition, following our own and others’ previous applications of PID to neuroscientific data 50.

Removed in the revised version: In a dynamical system such as the brain, one can calculate the amount of information flowing from the system’s past to its future, known as time-delayed mutual information (TDMI). Specifically, by denoting the past of variables as Xt-τ and Yt-τ and treating them as sources, and their joint future state (Xt, Yt), as target, one can apply the PID framework and decompose the information flowing from past to future as

Earlier version: Applying ΦID to this quantity allows us to distinguish between redundant, unique, and synergistic information shared with respect to the future variables Xt, Yt 49,50. Importantly, this framework, has identified Syn(Xt-r, Yt-r; Xt,Yt) with the capacity of the system to exhibit emergent behaviour 186 as well as a stronger notion of redundancy, in which information is shared by X and Y in both past and future. Accordingly, using the MMI-ΦID decomposition for Gaussian variables, we use

Revised version: In a dynamical system such as the brain, one can calculate the amount of information flowing from the system’s past to its future, known as time-delayed mutual information (TDMI). Specifically, by denoting the past of variables as Xt-τ and Yt-τ and treating them as sources, and their joint future state (Xt, Yt), as target, one can apply the PID framework and decompose the information flowing from past to future as Applying ΦID to this quantity allows us to distinguish between redundant, unique, and synergistic information shared with respect to the future variables Xt, Yt 49,50. Importantly, this framework, has identified,;, with the capacity of the system to exhibit emergent behaviour 192 as well as a stronger notion of redundancy, in which information is shared by X and Y in both past and future. Accordingly, using the MMI-ΦID decomposition for Gaussian variables, we use

Earlier version: Here, we used the Gaussian solver implemented in the JIDT toolbox 187 to obtain TDMI, synergy and redundancy between each pair of brain regions, based on their HRF-deconvolved BOLD signal timeseries” 49,50.

Revised version: Here, we used the Gaussian solver implemented in the JIDT toolbox 193 to obtain TDMI, synergy and redundancy between each pair of brain regions, based on their HRF-deconvolved BOLD signal timeseries” 49,50.

Revised measure of integrated information from Integrated Information Decomposition

Removed in the revised version: Through the framework of Integrated Information Decomposition, we can decompose the constituent elements of Φ, the formal measure of integrated information proposed by Integrated Information Theory to quantify consciousness 31. Note that several variants of Φ have been proposed over the years, including the original formulation of Tononi 29, other formulations based on causal perturbation 113,188 and others (see 48,189 for comparative reviews). Here, we focus on the “empirical Φ” measure of Seth and Barrett 190, based on the measures by Balduzzi and Tononi (2008) 31 and adapted to applications to experimental data. It is computed as

Earlier version: and it quantifies how much temporal information is contained in the system over and above the information in its past. This measure is easy to compute (compared with other Φ measures) 192 and represents a noteworthy attempt to formalise the powerful intuitions underlying IIT. However, once the original formulation from Balduzzi and Tononi is rendered suitable for practical empirical application 190,194 the resulting mathematical formulation has known shortcomings, including the fact that it can yield negative values in some cases - which are hard to interpret and seemingly paradoxical, as it does not seem plausible for a system to be “negatively integrated” or an organism to have negative consciousness 190,194.

Revised version: Through the framework of Integrated Information Decomposition, we can decompose the constituent elements of Φ, the formal measure of integrated information proposed by Integrated Information Theory to quantify consciousness 31. Note that several variants of Φ have been proposed over the years, including the original formulation of Tononi 29, other formulations based on causal perturbation 113,194 and others (see 48,195 for comparative reviews). Here, we focus on the “empirical Φ” measure of Seth and Barrett 196, based on the measures by Balduzzi and Tononi (2008) 31 and adapted to applications to experimental data. It is computed as and it quantifies how much temporal information is contained in the system over and above the information in its past. This measure is easy to compute (compared with other Φ measures) 197 and represents a noteworthy attempt to formalise the powerful intuitions underlying IIT. However, once the original formulation from Balduzzi and Tononi is rendered suitable for practical empirical application 196,198 the resulting mathematical formulation has known shortcomings, including the fact that it can yield negative values in some cases - which are hard to interpret and seemingly paradoxical, as it does not seem plausible for a system to be “negatively integrated” or an organism to have negative consciousness 196,198.

Interestingly, with ΦID it can be formally demonstrated 49 that Φ is composed of different information atoms: it contains all the synergistic information in the system, the unique information transferred from X to Y and vice versa, and, importantly, the subtraction of redundancy - which explains why Φ can be negative in redundancy-dominated systems.

Removed in the revised version: To address this fundamental shortcoming, Mediano, Rosas and colleagues 49 introduced a revised measure of integrated information, ΦR, which consists of the original Φ with the redundancy added back in:

Earlier version: where Red(X, Y) is defined in Eq. (4). This measure is computationally tractable and preserves the original intuition of integrated information as measuring the extent to which “the whole is greater than the sum of its parts”, since it captures only synergistic and transferred information. Crucially, thanks to Integrated Information Decomposition, it can be proved that the improved formulation of integrated information that we adopt here is guaranteed to be non-negative 49 - thereby avoiding a major conceptual limitation of the original formulation of Φ.

Revised version: To address this fundamental shortcoming, Mediano, Rosas and colleagues 49 introduced a revised measure of integrated information, ΦR, which consists of the original Φ with the redundancy added back in: where Red(X, Y) is defined in Eq. (4). This measure is computationally tractable and preserves the original intuition of integrated information as measuring the extent to which “the whole is greater than the sum of its parts”, since it captures only synergistic and transferred information. Crucially, thanks to Integrated Information Decomposition, it can be proved that the improved formulation of integrated information that we adopt here is guaranteed to be non-negative 49 - thereby avoiding a major conceptual limitation of the original formulation of Φ.

Earlier version: Note that the formula for ΦWMS above stems from what is known as IIT 2.0, but TDMI is by no means the only way of quantifying the dynamical structure of a system: indeed, subsequent developments in IIT 3.0 used alternative metrics with a more explicit focus on causal interpretations 113, which were in turn replaced in the latest iteration known as IIT 4.0 188,195. We do not consider the alternative measure of integrated information proposed in IIT 3.0 because it is computationally intractable for systems bigger than a small set of logic gates, and it is not universally well-defined 194.

Revised version: Note that the formula for ΦWMS above stems from what is known as IIT 2.0, but TDMI is by no means the only way of quantifying the dynamical structure of a system: indeed, subsequent developments in IIT 3.0 used alternative metrics with a more explicit focus on causal interpretations 113, which were in turn replaced in the latest iteration known as IIT 4.0 194,199. We do not consider the alternative measure of integrated information proposed in IIT 3.0 because it is computationally intractable for systems bigger than a small set of logic gates, and it is not universally well-defined 198.

Gradient of redundancy-to-synergy relative importance to identify the synergistic workspace

After building networks of synergistic and redundant interactions between each pair of regions of interest (ROIs), we determined the role of each ROI in terms of its relative engagement in synergistic or redundant interactions. Following the procedure previously described 50, we first calculated the nodal strength of each brain region as the sum of all its interactions in the group-averaged matrix (Figure S1). Then, we ranked all 454 regions based on their nodal strength (with higher-strength regions having higher ranks). This procedure was done separately for networks of synergy and redundancy. Subtracting each region’s redundancy rank from its synergy rank yielded a gradient from negative (i.e., ranking higher in terms of redundancy than synergy) to positive (i.e., having a synergy rank higher than the corresponding redundancy rank; note that the sign is arbitrary).

It is important to note that the gradient is based on relative - rather than absolute - differences between regional synergy and redundancy; consequently, a positive rank difference does not necessarily mean that the region’s synergy is greater than its redundancy; rather, it indicates that the balance between its synergy and redundancy relative to the rest of the brain is in favour of synergy - and vice versa for a negative gradient.

Subdivision of workspace nodes into gateways and broadcasters

Earlier version: To identify which regions within the workspace play the role of gateways or broadcasters postulated in our proposed architecture, we followed a procedure analogous to the one adopted to identify the gradient of redundancy-synergy relative importance, but replacing the node strength with the node participation coefficient. The participation coefficient Pi quantifies the degree of connection that a node entertains with nodes belonging to other modules: the more of a node’s connections are towards other modules, the higher its participation coefficient will be 60,196. Conversely, the participation coefficient of a node will be zero if its connections are all with nodes belonging to its own module.

Revised version: To identify which regions within the workspace play the role of gateways or broadcasters postulated in our proposed architecture, we followed a procedure analogous to the one adopted to identify the gradient of redundancy-synergy relative importance, but replacing the node strength with the node participation coefficient. The participation coefficient Pi quantifies the degree of connection that a node entertains with nodes belonging to other modules: the more of a node’s connections are towards other modules, the higher its participation coefficient will be 60,200. Conversely, the participation coefficient of a node will be zero if its connections are all with nodes belonging to its own module. Here, κis is the strength of positive connections between node i and other nodes in module s, kiis the strength of all its positive connections, and M is the number of modules in the network. The participation coefficient ranges between zero (no connections with other modules) and one (equal connections to all other modules) 60,200.

Removed in the revised version: Here, κis is the strength of positive connections between node i and other nodes in module s, ki is the strength of all its positive connections, and M is the number of modules in the network. The participation coefficient ranges between zero (no connections with other modules) and one (equal connections to all other modules) 60,196.

Here, modules were set to be the seven canonical resting-state networks identified by Yeo and colleagues 63, into which the Schaefer parcellation is already divided 62, with the addition of an eighth subcortical network comprising all ROIs of the Tian subcortical network 64. The brain’s RSNs were chosen as modules because of their distinct and well-established functional roles, which fit well with the notion of modules as segregated and specialised processing systems interfacing with the global workspace. Additionally, having the same definition of modules (i.e., RSNs) for synergy and redundancy allowed us to compute their respective participation coefficients in an unbiased way.

Separately for connectivity matrices of synergy and redundancy, the participation coefficient of each brain region was calculated. Then, regions belonging to the synergistic workspace were ranked, so that higher ranks indicated higher participation coefficient. Finally, the redundancy-based participation coefficient rank of each workspace region was subtracted from its corresponding synergy-based participation coefficient rank, to quantify – within the workspace – whether regions have relatively more diverse connectivity in terms of synergy, or in terms of redundancy.

This procedure yielded a gradient over workspace regions, from negative (i.e. having a more highly ranked participation coefficient based on redundancy than synergy) to positive (i.e. having a more highly ranked participation coefficient based on synergy than redundancy). Note that as before, the sign of this gradient is arbitrary, and it is based on relative rather than absolute difference. Workspace regions with a positive gradient value were classified as “gateways”, since they have synergistic interactions with many brain modules. In contrast, workspace regions with a negative value of the gradient - i.e. those whose redundancy rank is higher than their synergy rank, in terms of participation coefficient - were labelled as workspace “broadcasters”, since they possess information that is duplicated across multiple modules in the brain.

Statistical Analysis

Network Based Statistic

The network-based statistic approach 67 was used to investigate the statistical significance of propofol-induced or DOC-induced alterations. This nonparametric statistical method is designed to control the family-wise error due to multiple comparisons, for application to graph data. Connected components of the graph are identified from edges that survive an a-priori statistical threshold (F-contrast; here we set the threshold to an F-value of 9, two-sided, with an alpha level of 0.05). In turn, the statistical significance of such connected components is estimated by comparing their topology against a null distribution of the size of connected components obtained from non-parametric permutation testing. This approach rejects the null hypothesis on a component-by-component level, and therefore achieves superior power compared to mass-univariate approaches 67.

Testing for shared effects across datasets

Earlier version: We sought to detect changes that are common across datasets, to rule out possible propofol- or DOC-specific effects that are not related to consciousness per se 53. To this end, we employed a null hypothesis significance test under the composite null hypothesis that at least one dataset among those considered here has no effect. In other words, for the null hypothesis to be rejected we demand that all comparisons exhibit non-zero effects. As usual, the test proceeds by comparing an observed test statistic with a null distribution. The test statistic is the minimum of the three F-scores obtained in the comparisons of interest (DOC vs awake; anaesthesia vs awake; and anaesthesia vs recovery), and the null distribution is sampled by randomly reshuffling exactly one dataset (picked at random) at a time and recalculating the F-scores. By shuffling exactly one dataset (instead of all of them), we are comparing the observed data against the “least altered” version of the data that is still compatible with the null hypothesis. This is a type of least favourable configuration test 197, which is guaranteed to control the false positive rate below a set threshold (here, 0.05). The details of this test will be described in a future publication. Common changes across the three states of consciousness were then identified as edges (defined in terms of ΦR) that were either (i) increased in DOC compared with control; (ii) increased during anaesthesia compared with wakefulness; and (iii) increased during anaesthesia compared with post-anaesthetic recovery; or (i) decreased in DOC compared with control; (ii) decreased during anaesthesia compared with wakefulness; and (iii) decreased during anaesthesia compared with post-anaesthetic recovery.

Revised version: We sought to detect changes that are common across datasets, to rule out possible propofol-or DOC-specific effects that are not related to consciousness per se 53. To this end, we employed a null hypothesis significance test under the composite null hypothesis that at least one dataset among those considered here has no effect. In other words, for the null hypothesis to be rejected we demand that all comparisons exhibit non-zero effects. As usual, the test proceeds by comparing an observed test statistic with a null distribution. The test statistic is the minimum of the three F-scores obtained in the comparisons of interest (DOC vs awake; anaesthesia vs awake; and anaesthesia vs recovery), and the null distribution is sampled by randomly reshuffling exactly one dataset (picked at random) at a time and recalculating the F-scores. By shuffling exactly one dataset (instead of all of them), we are comparing the observed data against the “least altered” version of the data that is still compatible with the null hypothesis. This is a type of least favourable configuration test 201, which is guaranteed to control the false positive rate below a set threshold (here, 0.05). The details of this test will be described in a future publication. Common changes across the three states of consciousness were then identified as edges (defined in terms of ΦR) that were either (i) increased in DOC compared with control; (ii) increased during anaesthesia compared with wakefulness; and (iii) increased during anaesthesia compared with post-anaesthetic recovery; or (i) decreased in DOC compared with control; (ii) decreased during anaesthesia compared with wakefulness; and (iii) decreased during anaesthesia compared with post-anaesthetic recovery.

Spatial autocorrelation-preserving null model for correlation

Earlier version: The significance of correlation between nodes’ participation coefficient based on different definitions of modules (a-priori as resting-state networks or in a data-driven fashion from Louvain community detection) was assessed using a spatial permutation test which generates a null distribution of 10,000 randomly rotated brain maps with preserved spatial covariance (‘spin test’), to ensure robustness to the potential confounding effects of spatial autocorrelation 198–200.

Revised version: The significance of correlation between nodes’ participation coefficient based on different definitions of modules (a-priori as resting-state networks or in a data-driven fashion from Louvain community detection) was assessed using a spatial permutation test which generates a null distribution of 10,000 randomly rotated brain maps with preserved spatial covariance (‘spin test’), to ensure robustness to the potential confounding effects of spatial autocorrelation 202–204.

Data Availability

The raw data analysed during the current study are available on request from the following authors. Propofol anaesthesia, Disorders of Consciousness and test-retest datasets: Dr. Emmanuel A. Stamatakis (University of Cambridge, Division of Anaesthesia; email: eas46@cam.ac.uk). The Human Connectome Project datasets are freely available from http://www.humanconnectome.org/.

Code availability

The Java Information Dynamics Toolbox v1.5 is freely available online: (https://github.com/jlizier/jidt). The CONN toolbox version 17f is freely available online (http://www.nitrc.org/projects/conn). The Brain Connectivity Toolbox code used for graph-theoretical analyses is freely available online (https://sites.google.com/site/bctnet/). The HRF deconvolution toolbox v2.2 is freely available online: (https://www.nitrc.org/projects/rshrf). The code for spin-based permutation testing of cortical correlations is freely available at https://github.com/frantisekvasa/rotate_parcellation. We have made freely available MATLAB/Octave and Python code to compute measures of Integrated Information Decomposition of timeseries with the Gaussian MMI solver, at https://github.com/Imperial-MIND-lab/integrated-info-decomp.

Removed in the revised version: Supporting information